```python
# Initialize Otter
import otter
grader = otter.Notebook("Lab_7_optimization.ipynb")
```

# Setting up a function for fmin

Set up the function that we will be passing to fmin

Slides: https://docs.google.com/presentation/d/1m-8UphLgGLVaDKKcZ-NA_ejgOFlkD1Rlz3_sW0fHBM8/edit?usp=sharing

```python
# The usual imports
import numpy as np
import matplotlib.pyplot as plt

```

```python
# Imports from week 6
# This bit of code tells Jupyter notebooks to also look in the week_6_optimization folder for files
import os
import sys
sys.path.insert(0, os.path.abspath('../Week_6_matrices'))

# Routines from HW 3, Lab 6
# These are the routines used in the lab and/or you need to fill in for the homework (lab 6)
from arm_routines import create_arm_geometry, set_angles_of_arm_geometry, get_gripper_location
```

```python
# These are the routines used in the lab and/or you need to fill in for the homework (lab 7, 2nd half of homework)
from arm_optimization_routines import vector_to_goal, distance_to_goal, func_for_fmin, angles_list_to_numpy
```

```python
# These commands will force JN to actually re-load the external file when you re-execute the import command
%load_ext autoreload
%autoreload 2
```

Output:
```
The autoreload extension is already loaded. To reload it, use:
  %reload_ext autoreload
```

## Distance from gripper location to target

TODO: Edit vector_to_goal and distance_to_goal in arm_optimization_routines.py

```python
# Create the arm geometry
base_size_param = (1.0, 0.5)
link_sizes_param = [(0.5, 0.25), (0.3, 0.1), (0.2, 0.05)]
palm_width_param = 0.1
finger_size_param = (0.075, 0.025)

arm_geometry = create_arm_geometry(base_size_param, link_sizes_param, palm_width_param, finger_size_param)

# Set the angles of the arm
angles_gripper_check = [np.pi/6.0, -np.pi/4, 1.5 * np.pi/4, [np.pi/3.0, -np.pi/8.0, np.pi/6.0]]
set_angles_of_arm_geometry(arm_geometry, angles_gripper_check)
```

```python
mat = np.array([1.0,0.0])
rot_mat = np.array([0.5,0.5])
rotated_mat = mat@rot_mat
print(rotated_mat)
```

Output:
```
0.5
```

```python
# Get the gripper location
x, y = get_gripper_location(arm_geometry)
print(f"Gripper location ({x}, {y}), should be (-0.383, 1.323)")
```

Output:
```
Gripper location (-0.38299317823125056, 1.3230167925101477), should be (-0.383, 1.323)
```

```python
target = np.array([0.75, 1.25])
vx, vy = vector_to_goal(arm_geometry, target)
print(f"Vector to goal ({vx}, {vy}), should be (1.133, -0.073)")
```

Output:
```
Vector to goal (1.1329931782312506, -0.0730167925101477), should be (1.133, -0.073)
```

```python
d_to_target = distance_to_goal(arm_geometry, target)
print(f"Distance to goal {d_to_target}, should be 1.1353")
```

Output:
```
Distance to goal 1.1353435576542548, should be 1.1353
```

```python
grader.check("dist_func")
```

Output:
```
dist_func results: All test cases passed!
```

# Edit the function that will be passed to fmin

TODO: edit **func_for_fmin** to take in a numpy array of angles and return the distance

```python
# This is a helper function to convert the list of angles in angles_gripper into a numpy array, which is
#  what fmin works with
angles_numpy = angles_list_to_numpy(angles_gripper_check)
check_dist = func_for_fmin(angles_numpy, arm_geometry, target)
print(f"Check func for min {check_dist}, should be 1.1353")

```

Output:
```
Check func for min 1.1353435576542548, should be 1.1353
```

```python
grader.check("fmin_func")
```

Output:
```
fmin_func results: All test cases passed!
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"none"}
# List of URLS TCW3 (creates a set)
websites = {"none"}
# Approximate number of hours, including lab/in-class time
hours = 7
```

```python
grader.check("hours_collaborators")
```

Output:
```
hours_collaborators results: All test cases passed!
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit through gradescope, Lab 7 optimization. Don't forget arm_optimization_routines.py (we will supply data files and matrix/object routines). If your arm_routines.py does not work, leave it out and the autograder will use our solution

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Output:
```
Running your submission against local test cases...


File c:\Users\yeasshhhh\anaconda3\Lib\site-packages\otter\check\notebook.py:505, in Notebook.export.<locals>.continue_export()
    503 if run_tests:
    504     print("Running your submission against local test cases...\n")
--> 505     results = grade_zip_file(zip_path, nb_path, self._tests_dir)
    506     print(
    507         "Your submission received the following results when run against " + \
    508         "available test cases:\n\n" + indent(results.summary(), "    "))
    510 if display_link:
    511     # create and display output HTML

File c:\Users\yeasshhhh\anaconda3\Lib\site-packages\otter\check\utils.py:103, in grade_zip_file(zip_path, nb_arcname, tests_dir)
    100 print(results.stdout.decode("utf-8"))
    102 if results.stderr:
--> 103     raise RuntimeError(results.stderr.decode("utf-8"))
    105 with open(results_path, "rb") as f:
    106     results = dill.load(f)

RuntimeError: c:\Users\yeasshhhh\anaconda3\Lib
```

Tests for dist_func (points: 4)
```
>>> assert(np.isclose(vx, 1.133, atol=0.01))
```
```
>>> assert(np.isclose(vy, -0.073, atol=0.01))
```
```
>>> assert(np.isclose(d_to_target, 1.1353, atol=0.01))
```


Tests for fmin_func (points: 5)
```
>>> assert(np.isclose(check_dist, 1.1353, atol=0.1))
```


Tests for hours_collaborators (points: 1)
```
>>> assert(not "not filled out" in worked_with_names)
```
```
>>> assert(not "not filled out" in websites)
```
```
>>> assert(hours > 0)
```
