```python
# Initialize Otter
import otter
grader = otter.Notebook("lec_act_8_calculus.ipynb")
```

# Lecture activity: calculus 

Practice calculus with this function:

f(x) = (3/x) ln(x) - x^2 / 10.0 - x / 3.0 + (x-3) * x * x

# Create the function

Write the Python function for the above equation

```python
# Do the imports
import numpy as np
import matplotlib.pyplot as plt
from scipy.integrate import quad
from scipy.optimize import fminbound

```

```python
# Use my_func for the name of the Python function
def my_func(x):
    return (3/x)*np.log(x)-x**2/10-x/3+(x-3)*x*x
```

```python
print(f"Func at 1.0 is {my_func(1.0)}, expected -2.433333")
print(f"Func at 2.0 is {my_func(2.0)}, expected -4.02694589")
```

Output:
```
Func at 1.0 is -2.4333333333333336, expected -2.433333
Func at 2.0 is -4.026945895826749, expected -4.02694589
```

```python
grader.check("create_func")
```

Output:
```
create_func results: All test cases passed!
```

# Plot the function, along with the bounds and an x at minima and maxima

Use 0.5 and 3.0 as the integration bounds

```python
# Integration limits
a = 0.5
b = 3.0
xs = np.linspace(a, b, num=100)
ys = my_func(xs)
```

```python
# Find the x locations of the two bends
#
first_max = fminbound(lambda x: -my_func(x), a, b)  
first_min = fminbound(my_func, a, b)
# END PROMPT """

print(f"Max {first_max}, expected 0.94813, Min {first_min}, expected 2.091114")
```

Output:
```
Max 0.9481378678543182, expected 0.94813, Min 2.091115061958074, expected 2.091114
```

```python
grader.check("plot_fn")
```

Output:
```
plot_fn results: All test cases passed!
```

# Derivative
TODO Numerically alculate the derivatives at the first_max/first_min points (what should the derivative be?)

Pick a reasonable **h** for the derivative calculation. You can do forward, backward, or central

```python
def my_deriv(x, h = 1e-6):
    deriv =  (my_func(x+h) - my_func(x-h)) / (2.0 * h)
    return deriv
```

```python
h = 1e-6
x = first_max
central_deriv = (my_func(x+h) - my_func(x-h)) / (2.0 * h)
print(f"Central deriv at max {central_deriv}, expected 0")
deriv_max = my_deriv(first_max)
deriv_min = my_deriv(first_min)

print(f"Deriv at max {deriv_max}, expected 0, Deriv at min {deriv_min}, expected 0")    

```

Output:
```
Central deriv at max -1.7541523789077473e-07, expected 0
Deriv at max -1.7541523789077473e-07, expected 0, Deriv at min -3.455014052633487e-06, expected 0
```

```python
grader.check("derivative")
```

Output:
```
derivative results: All test cases passed!
```

# Integration

TODO Calculate the area under the curve using trapz or quad, whichever you perfer.

```python

area_2 = np.trapz(ys, xs)
print(area_2)
area_under_curve = quad(my_func,a,b)
print(area_under_curve)
# END PROMPT"""
```

Output:
```
-7.905557521920043
(-7.905047746325095, 9.901345449853068e-11)
```

```python
grader.check("area")
```

Output:
```
area results: All test cases passed!
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"none"}
# List of URLS TCW3 (creates a set)
websites = {"<https://stackoverflow.com/questions/10146924/finding-the-maximum-of-a-function>"}
# Approximate number of hours, including lab/in-class time
hours = 1.25
```

```python
grader.check("hours_collaborators")
```

Output:
```
hours_collaborators results: All test cases passed!
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit through gradescope, Lecture activity 8, Calculus

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Output:
```
Running your submission against local test cases...

File c:\Users\yeasshhhh\anaconda3\Lib\threading.py:1112, in Thread.join(self, timeout)
   1109     raise RuntimeError("cannot join current thread")
   1111 if timeout is None:
-> 1112     self._wait_for_tstate_lock()
   1113 else:
   1114     # the behavior of a negative timeout isn't documented, but
   1115     # historically .join(timeout=x) for x<0 has acted as if timeout=0
   1116     self._wait_for_tstate_lock(timeout=max(timeout, 0))

File c:\Users\yeasshhhh\anaconda3\Lib\threading.py:1132, in Thread._wait_for_tstate_lock(self, block, timeout)
   1129     return
   1131 try:
-> 1132     if lock.acquire(block, timeout):
   1133         lock.release()
   1134         self._stop()

KeyboardInterrupt:
```

Tests for area (points: 1)
```
>>> assert(np.isclose(area_under_curve[0], -7.90504, atol=0.001))
```


Tests for create_func (points: 1)
```
>>> assert(np.isclose(my_func(1.0), -2.433333, atol=0.001))
```
```
>>> assert(np.isclose(my_func(2.0), -4.02694589, atol=0.001))
```


Tests for derivative (points: 1)
```
>>> assert(np.isclose(deriv_max, 0.0, atol=0.0001))
```
```
>>> assert(np.isclose(deriv_min, 0.0, atol=0.0001))
```


Tests for hours_collaborators (points: 1)
```
>>> assert(not "not filled out" in worked_with_names)
```
```
>>> assert(not "not filled out" in websites)
```
```
>>> assert(hours > 0)
```


Tests for plot_fn (points: 1)
```
>>> assert(np.isclose(first_max, 0.94813, atol=0.1))
```
```
>>> assert(np.isclose(first_min, 2.091114, atol=0.1))
```
