```python
# Initialize Otter
import otter
grader = otter.Notebook("lec_act_5_predator_prey.ipynb")
```

#  Predator-prey iterative functions 

We're going to start implementing a simple version of the predator/prey relationship, a classic differential equations problem also known as the Lotka-Volterra equations. You don't need to know a whole lot about them, except that there are two variables to track - prey and predator - and that how the numbers of prey/predator changes depends on both the current prey values AND the predator values (the differential equations part)

Resources:
-   https://en.wikipedia.org/wiki/Lotka%E2%80%93Volterra_equations for a more general theoretical introduction
-   https://www.kristakingmath.com/blog/predator-prey-systems a better "what is it" description
-   https://scientific-python.readthedocs.io/en/latest/notebooks_rst/3_Ordinary_Differential_Equations/02_Examples/Lotka_Volterra_model.html - implementation in sci-py (which is what you should really use for this type of problem - scipy's ode solver)

Also see the slides for the lab: https://docs.google.com/presentation/d/1wd1SpTJiezfroDizaFkA6UxCgMQE_A4ZKCyJoNfa0cM/edit?usp=sharing

In this lecture activity: You're going to 
- practice writing a function from a function description 
- practice turning a bit of math into an iterative function
- practice calling functions from functions


Note: You can write the entire solver in one function. But that's a) rather hard to debug and b) can result in making the most common mistake with iterative functions - not computing the new values from the old ones.

```python
# Doing the imports for you
import numpy as np
import matplotlib.pyplot as plt
```

## First function - compute the new prey value from the prey and predator

Function input: 
- Current prey value (number)
- current predator value (number)
- the parameters as a dictionary (see calling function code)

Output
- New prey value (number)

Equation:
- dprey/dt = "Prey reproduce" * prey - "Prey eaten" * prey * predator
- prey = prey + delta_t * dprey/dt

Function name: Use **compute_prey_from_prey_and_predator**

```python

# TODO: Fill in the parameters (see description above). Don't forget to comment what the input paramters are
def compute_prey_from_prey_and_predator(prey, predator, params):
    # TODO: Calculate the new prey value from the input prey/predator values and delta t (see equation in the cell above)
    #  Note: To get, eg, the "Prey reproduce" value use params["Prey reproduce"].
    
    dprey_dt = params["Prey reproduce"] * prey - params["Prey eaten"] * prey * predator
    new_prey =  prey + params["delta t"] * dprey_dt
    return new_prey
    pass
```

```python
# Test code
# Tie the number of time steps to the total number of days
delta_t = 0.1
n_days = 40
n_time_steps = int(n_days / delta_t)
# Store everything you need to compute the system in the one dictionary
params = {"Prey reproduce":1.0,
          "Prey eaten":0.02,
          "Predator loss":1.2,
          "Predator reproduce":0.03,
          "delta t": delta_t,   # unit: days
          "n days": n_days,     # unit: days
          "n time steps": n_time_steps}

prey_initial = 100
predator_initial = 100

prey_new = compute_prey_from_prey_and_predator(prey=prey_initial, predator=predator_initial, params=params)

print(f"Checking prey new {prey_new}, should be 90")
```

Output:
```
Checking prey new 90.0, should be 90
```

```python
grader.check("prey_from_prey_and_predator")
```

Output:
```
prey_from_prey_and_predator results: All test cases passed!
```

## Second function (predator)

Compute the new predator value from the prey and predator.

Input: 
- Current prey value (number)
- current predator value (numer)
- the parameters as a dictionary (see calling function)

Output: 
- New predator value (number)

Equation:
- dpredator/dt = - "Predator loss" * predator + "Predator reproduce" * prey * predator
- predator = predator + delta_t * dpredator/dt

Function name

Use **compute_predator_from_prey_and_predator**


```python

# TODO: Fill in the parameters (see description above). Don't forget to comment what the input paramters are
def compute_predator_from_prey_and_predator(prey , predator, params):
    # TODO: Calculate the new predator value from the input prey/predator values and delta t (see equation)
    #  Note: To get, eg, the "Prey reproduce" value use params["Prey reproduce"].

    dpredattor_dt = -params["Predator loss"] * predator + params["Predator reproduce"] * prey * predator
    new_predator = predator +delta_t * dpredattor_dt
    return new_predator
    pass
```

```python
# Check code - uses parameters defined in question 1
predator_new = compute_predator_from_prey_and_predator(prey=prey_initial, predator=predator_initial, params=params)

print(f"Checking predator new {predator_new}, should be 118")
```

Output:
```
Checking predator new 118.0, should be 118
```

```python
grader.check("predator_from_prey_and_predator")
```

Output:
```
predator_from_prey_and_predator results: All test cases passed!
```

# Third function (call both)

Put the two functions together.

Input: 
- Current prey value (number)
- current predator value (numer), the parameters as a dictionary (see calling function)

Output: New prey, predator values (tuple)

Functionality: Should just call the two functions, one after the other

Function name: use **compute_one_time_step**

TODO: Once you have this working correctly make one small change - use the new prey value you calculate in the call to calculate the predator value (instead of the input prey value). How much is the result off by? This is a really, really common error and one that is difficult to catch. By writing the code this way (two functions, then calling one function after the other) you're less likely to make that mistake in the first place

Don't forget to put it back to the correct answer

```python

# TODO: Fill in the parameters (see description above). Don't forget to comment what the input paramters are
def compute_one_time_step(prey, predator, params):
    # TODO: Calculate the new prey/predator values from the input prey/predator values and delta t (see equation)
    # Do NOT re-write the equations - call the functions you already wrote
    new_prey = compute_prey_from_prey_and_predator(prey, predator, params)
    new_predator = compute_predator_from_prey_and_predator(prey, predator, params)
    return new_prey, new_predator 
    pass
```

```python
# Test code
prey_new2, predator_new2 = compute_one_time_step(prey_initial, predator_initial, params)

print(f"Checking prey new {prey_new2}, should be 90 predator new {predator_new2}, should be 118")
```

Output:
```
Checking prey new 90.0, should be 90 predator new 118.0, should be 118
```

```python
grader.check("compute_one_time_step")
```

Output:
```
compute_one_time_step results: All test cases passed!
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"none"}
# List of URLS TCW3 (creates a set)
websites = {"none"}
# Approximate number of hours, including lab/in-class time
hours = .75
```

```python
grader.check("hours_collaborators")
```

Output:
```
hours_collaborators results: All test cases passed!
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit through gradescope, Lecture activity week 5, predator prey.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Output:
```
Running your submission against local test cases...


File c:\Users\yeasshhhh\anaconda3\Lib\site-packages\otter\check\notebook.py:505, in Notebook.export.<locals>.continue_export()
    503 if run_tests:
    504     print("Running your submission against local test cases...\n")
--> 505     results = grade_zip_file(zip_path, nb_path, self._tests_dir)
    506     print(
    507         "Your submission received the following results when run against " + \
    508         "available test cases:\n\n" + indent(results.summary(), "    "))
    510 if display_link:
    511     # create and display output HTML

File c:\Users\yeasshhhh\anaconda3\Lib\site-packages\otter\check\utils.py:103, in grade_zip_file(zip_path, nb_arcname, tests_dir)
    100 print(results.stdout.decode("utf-8"))
    102 if results.stderr:
--> 103     raise RuntimeError(results.stderr.decode("utf-8"))
    105 with open(results_path, "rb") as f:
    106     results = dill.load(f)

RuntimeError: c:\Users\yeasshhhh\anaconda3\Lib
```

Tests for compute_one_time_step (points: 1)
```
>>> assert(np.isclose(prey_new, prey_new2))
```
```
>>> assert(np.isclose(predator_new, predator_new2))
```


Tests for hours_collaborators (points: 1)
```
>>> assert(not "not filled out" in worked_with_names)
```
```
>>> assert(not "not filled out" in websites)
```
```
>>> assert(hours > 0)
```


Tests for predator_from_prey_and_predator (points: 2)
```
>>> assert(np.isclose(predator_new, 118))
```


Tests for prey_from_prey_and_predator (points: 2)
```
>>> assert(np.isclose(prey_new, 90))
```
