```python
# Initialize Otter
import otter
grader = otter.Notebook("lec_act_2_functions.ipynb")
```

# Lecture goals

1. (more) Dictionaries for data encapsulation
2. Functions for functionality encapsulation

Lab slides: https://docs.google.com/presentation/d/1ykwwcQ0onMvAjUxfJmKl9tbo-rJPdB5pRwDEmpDsd-g/edit?usp=sharing

## Functions: 
Functions enable encapsulation of, well, functionality.

They're also a useful mental tool for organizing and structuring your thoughts on how to solve a given problem
1. Clearly define a bit of code that takes in some inputs, does some computation, then outputs some data
2. Makes it easier to test that code with different inputs
3. Practicalities: Prevents one of the most common sources of errors - re-using variable names

It's almost never wrong to encapsulate a bit of code in a function. It can slow down (a tiny bit) computation time, but can greatly reduce debugging time, so it's usually worth it.

Python's function syntax is beautifully designed to make it easy to set default values for parameters and pass back as much data as you want. We'll see more of that later; for this assignment we'll use the power of dictionaries to pass back "labeled" data.

In this lecture activity you're essentially going to copy over the code you did in the previous lecture activity, right-shift it to place it in the function, replace the **test_** variable name with the input parameter to the function, then add a return to return the answer. This is a pretty common approach for turning code into a function. 


```python
# Access all numpy functions as np.
import numpy as np
```

## Question 1: Stats on a list

### Calculate stats on a list

TODO: in calc_stats_from_list function 
- Calculate the mean of the negative and positive values
- Count the total number of negative/positive values
- Store the values in a dictionary

This function calculates the given stats from the list that is passed in. There is test code below this function.

TODO: 
 - Step 1 - copy your code from lecture activity 1 into the function. Right shift it so it is indented properly
 - Step 2 - change **test_list_one** to be **in_list** - the input parameter of the function
 - Step 3 - return the dictionary **dict_save_stats**

```python
def calc_stats_from_list(in_list):
    """ Calculate mean of positive numbers, mean of negatives numbers
    Separate the list into positive and negative numbers. Calculate the mean of each. Return those means, along with
     how many positive/negative numbers there were
    @param in_list : any list type
    @return - A dictionary with the desired stats"""

    # These are the stats we're calculating. This is more elegant/useful than creating four variables - it keeps all
    #  of the values in the same place and assigns a meaningful label (key) to them
    dict_save_stats = {"Mean positive": 0.0, "Mean negative": -0.0, "Count positive": 0, "Count negative": 0}

    # TODO: 
    #   Copy your code from lecture activity 1 here. Don't forget to change the name of test_list_one to be in_list
    count_negative=0
    count_positive=0
    neg_list=[]
    pos_list=[]

    for item in in_list:
        if item < 0:
            count_negative = count_negative+1
            neg_list.append(item)
        elif item > 0:
            count_positive = count_positive+1
            pos_list.append(item)

   
    mean_positive=sum(pos_list)/len(pos_list)

    mean_negative=sum(neg_list)/len(neg_list)

    dict_save_stats["Count negative"]=count_negative
    dict_save_stats["Count positive"]=count_positive
    dict_save_stats["Mean negative"]=mean_negative
    dict_save_stats["Mean positive"]=mean_positive

    # TODO Do the return here
    
    return dict_save_stats

```

### Test code for list

Create the arrays and test your function using those test arrays. Here's another advantage of functions - you can create test data for yourself to make sure the code is working right. Encapsulating the code in a function means you don't 
1. Accidentally change the code when switching from the test data to the real data
2. You can make more than one test 
3. You can run the tests more than once/all the time to double check that you didn't "break" the code

TODO: 
- Fill in the calc_stats_from_list function above
- Run the cell below - it will print out if your values are incorrect

Note that, below, we'll test this code one last time with randomly generated data

```python
# Test data
test_list_one = [-0.75, -0.25, 1.0 / 3.0, 2.0 / 3.0, 3.0 / 3.0]
test_list_res = calc_stats_from_list(test_list_one)

b_tests_passed = True
if not np.isclose(test_list_res["Mean positive"], 2.0 / 3.0):
    b_tests_passed = False
    print(f"Mean positive is not correct, should be {2.0/3.0}, got {test_list_res['Mean positive']}")

if not np.isclose(test_list_res["Mean negative"], -0.5):
    b_tests_passed = False
    print(f"Mean negative is not correct, should be -0.5, got {test_list_res['Mean negative']}")

if test_list_res["Count positive"] != 3:
    b_tests_passed = False
    print(f"Count positive numbers, should be 3, got {test_list_res['Count positive']}")

if test_list_res["Count negative"] != 2:
    b_tests_passed = False
    print(f"Count positive numbers, should be 2, got {test_list_res['Count negative']}")

if b_tests_passed:
    print("All array tests passed!")
```

Output:
```
All array tests passed!
```

```python
grader.check("list")
```

Output:
```
list results: All test cases passed!
```

# Question 2: Doing it again with a numpy array

## Fill in calc_stats_from_nparray

For this function, assume the input is a numpy array.

TODO: Same as the previous question, but this time copy in the numpy array code you wrote in lecture activity 1. Again:
- NO **if** statements or **for** loops - do this all with numpy operations

As before, test code is below

```python
def calc_stats_from_nparray(in_nparray):
    """ Calculate mean of positive numbers, mean of negatives numbers
    Separate the list into positive and negative numbers. Calculate the mean of each. Return those means, along with
     how many positive/negative numbers there were
    @param in_nparray : numpy array
    @return - A dictionary with the desired stats"""

    # TODO: Copy in the code, change the numpy array name to be the input to this function, and return the dictionary
    my_array = np.array(in_nparray)
    negs = my_array<0
    pos = my_array>0

    all_negs = my_array[negs]
    all_pos = my_array[pos]

    dict_save_stats_np = {}

    dict_save_stats_np["Count negative"] = np.count_nonzero(all_negs)
    dict_save_stats_np["Count positive"] = np.count_nonzero(all_pos)
    dict_save_stats_np["Mean negative"] = np.mean(all_negs)
    dict_save_stats_np["Mean positive"] = np.mean(all_pos)
    return dict_save_stats_np

```

### Test code for numpy array

This will print out if your function above is returning incorrect values

```python
test_nparray_one = np.array(test_list_one)  # Convert the previous test list to a numpy array
test_list_res = calc_stats_from_nparray(test_nparray_one)

b_tests_passed = True
if not np.isclose(test_list_res["Mean positive"], 2.0 / 3.0):
    b_tests_passed = False
    print(f"Mean positive is not correct, should be {2.0/3.0}, got {test_list_res['Mean positive']}")

if not np.isclose(test_list_res["Mean negative"], -0.5):
    b_tests_passed = False
    print(f"Mean negative is not correct, should be -0.5, got {test_list_res['Mean negative']}")

if test_list_res["Count positive"] != 3:
    b_tests_passed = False
    print(f"Count positive numbers, should be 3, got {test_list_res['Count positive']}")

if test_list_res["Count negative"] != 2:
    b_tests_passed = False
    print(f"Count positive numbers, should be 2, got {test_list_res['Count negative']}")

if b_tests_passed:
    print("All numpy array tests passed!")
```

Output:
```
All numpy array tests passed!
```

```python
grader.check("nparray")
```

Output:
```
nparray results: All test cases passed!
```

### Encapsulating the data slice 

In this problem we'll encapsulate the fancy slice you did in lab 2 where you get out the x,y,z columns for the wrist torque data. 

```python
# Load in the data
try:
    fname = "Data/proxy_pick_data.csv"
    pick_data = np.loadtxt(fname, dtype="float", delimiter=",")
except FileNotFoundError:
    print(f"File not found {fname}")


# TODO - copy over your code from Lab 1 here
pick_channel_data =pick_data[:,:-1]
print(pick_channel_data.shape)
```

Output:
```
(660, 1320)
```

```python

```

```python
def get_channel_data(all_data, n_picks, start_index, n_time_steps, n_total_dims, n_dims):
    """ Get the data for just one channel (eg, wrist torque)
    @param all_data - the pick_channel_data numpy array
    @param n_picks - number of picks (number of rows in all_data)
    @param start_index - where to start getting data from 
    @param n_time_steps - number of time steps
    @param n_total_dims - what the skip value is - the total number of channels
    @param n_dims - total number of dimensions to use (1, 2, or 3)
    @return Return array should be n_picks X (n_timesteps * n_dims)"""

    # TODO Your slice code goes here. Note that I kept most of the variable names the same, so you should only have
    #  to change the wrist torque specific ones
   
    wrist_torque_data = np.zeros((n_picks, n_time_steps * n_dims))
    wrist_torque_data[0:n_picks,0:n_time_steps] = pick_channel_data[:, start_index::n_total_dims]
    wrist_torque_data[:, 1::3] = pick_channel_data[:, start_index+1::n_total_dims]
    wrist_torque_data[:, 2::3] = pick_channel_data[:, start_index+2::n_total_dims]


    print(wrist_torque_data.shape)
    print(wrist_torque_data[0,-1])
    return wrist_torque_data
```

```python

```

```python
# Just for this problem I'm "hard-wiring" in these values - this is because I want to test the function
# with known values
n_picks = 660
wt_start_index = 3  # Wrist torque starts at 3
n_time_steps = 40
n_total_dims = 33

# Note - I know the wrist torque data has 3 dimensions (x,y,z)
wrist_torque_data = get_channel_data(pick_channel_data, 
                                     n_picks=n_picks, 
                                     start_index=wt_start_index,
                                     n_time_steps=n_time_steps,
                                     n_total_dims=n_total_dims,
                                     n_dims=3)
```

Output:
```
(660, 120)
-0.077284199
```

```python
print(pick_channel_data[0,-34:-1])
```

Output:
```
[ 2.69000015e+01 -4.37332067e-01  4.80904677e+00 -5.57173931e+00
 -4.33966910e-01 -1.72362242e-01 -7.72841990e-02 -2.62883353e+00
 -9.73721790e+00 -1.24977326e+00 -3.15093994e+00 -1.06048584e+00
 -7.78198242e-01  4.29791992e+02  0.00000000e+00  3.22799988e+01
 -1.61369193e+00 -3.28723884e+00  9.39963531e+00 -1.75476074e-01
  7.32421875e-01  3.81469730e-02  2.59864014e+02  0.00000000e+00
  2.95900002e+01 -2.54503632e+00 -9.80904388e+00 -1.22343707e+00
 -3.77655029e+00 -1.12915039e+00  7.01904297e-01  2.09880005e+02
  0.00000000e+00]
```

```python
grader.check("data_function")
```

Output:
```
data_function results: All test cases passed!
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"none"}
# List of URLS TCW3 (creates a set)
websites = {"none"}
# Approximate number of hours, including lab/in-class time
hours = 2.5
```

```python
grader.check("hours_collaborators")
```

Output:
```
hours_collaborators results: All test cases passed!
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit through gradescope, lecture activity 2 Functions. Be sure to read the info on the autograder before submitting.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Output:
```
Running your submission against local test cases...


File c:\Users\yeasshhhh\anaconda3\Lib\site-packages\otter\check\notebook.py:505, in Notebook.export.<locals>.continue_export()
    503 if run_tests:
    504     print("Running your submission against local test cases...\n")
--> 505     results = grade_zip_file(zip_path, nb_path, self._tests_dir)
    506     print(
    507         "Your submission received the following results when run against " + \
    508         "available test cases:\n\n" + indent(results.summary(), "    "))
    510 if display_link:
    511     # create and display output HTML

File c:\Users\yeasshhhh\anaconda3\Lib\site-packages\otter\check\utils.py:103, in grade_zip_file(zip_path, nb_arcname, tests_dir)
    100 print(results.stdout.decode("utf-8"))
    102 if results.stderr:
--> 103     raise RuntimeError(results.stderr.decode("utf-8"))
    105 with open(results_path, "rb") as f:
    106     results = dill.load(f)

RuntimeError: c:\Users\yeasshhhh\anaconda3\Lib
```

Tests for data_function (points: 2)
```
>>> assert(wrist_torque_data.shape[0] == pick_data.shape[0])
```
```
>>> assert(wrist_torque_data.shape[1] == 3 * 40)
```
```
>>> assert(np.isclose(wrist_torque_data[0, 0], -0.268183058, atol=0.001))
```
```
>>> assert(np.isclose(wrist_torque_data[0, -1], -0.-0.077284199, atol=0.001))
```


Tests for hours_collaborators (points: 1)
```
>>> assert(not "not filled out" in worked_with_names)
```
```
>>> assert(not "not filled out" in websites)
```
```
>>> assert(hours > 0)
```


Tests for list (points: 1)
```
>>> res = calc_stats_from_list([-1, 2, -3, 4, -5])
>>> assert np.isclose(res["Mean positive"], 3.0)
```
```
>>> res = calc_stats_from_list([-1, 2, -3, 4, -5])
>>> assert np.isclose(res["Mean negative"], -3.0)
```
```
>>> res = calc_stats_from_list([-1, 2, -3, 4, -5])
>>> assert np.isclose(res["Count positive"], 2)
```
```
>>> res = calc_stats_from_list([-1, 2, -3, 4, -5])
>>> assert np.isclose(res["Count negative"], 3)
```


Tests for nparray (points: 1)
```
>>> test_data = np.array([-1, 2, -3, 4, -5])
>>> res = calc_stats_from_nparray(test_data)
>>> assert np.isclose(res["Mean positive"], 3.0)
```
```
>>> test_data = np.array([-1, 2, -3, 4, -5])
>>> res = calc_stats_from_nparray(test_data)
>>> assert np.isclose(res["Mean negative"], -3.0)
```
```
>>> test_data = np.array([-1, 2, -3, 4, -5])
>>> res = calc_stats_from_nparray(test_data)
>>> assert np.isclose(res["Count positive"], 2)
```
```
>>> test_data = np.array([-1, 2, -3, 4, -5])
>>> res = calc_stats_from_nparray(test_data)
>>> assert np.isclose(res["Count negative"], 3)
```
